--- 
title_meta  : Capitolo 2
title       : Vettori
description : "In questo corso ti porteremo a Las Vegas, dove imparerai ad analizzare in R il risultato delle tue scommesse utilizzando i vettori! Dopo aver completato questo capitolo riuscirai a creare vettori in R, assegnare loro nomi, selezionare elementi e comparare differenti vettori tra loro."

--- type:NormalExercise xp:100 skills:1 key:7329436201
## Crea un vettore

Ti senti fortunato? Speriamo di si siccome questo capitolo ti porta a Las Vegas, conosciuta anche come il paradiso degli statistici!

Grazie a R e alle tue nuove abilità nell'analizzare dati, imparerai a migliorare le tue performance al tavolo da gioco e avviare la tua carriera di scommettitore professionista. Questo capitolo ti mostrerà come tenere sotto controllo i risultati ottenuti e come effettuare alcune analisi ex-post. Prossima fermata: LAS VEGAS!!

*** =instructions
- Ti ricordi quanto imparato nel capitolo 1? Prova ad assegnare il valore `"Go!"` alla variabile `vegas`. Ricorda: R è case sensitive!

*** =hint
Scrivi la seguente riga di codice nell'editor:
```
vegas <- "Go!"
```

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# Definisci la variabile vegas
vegas <- 
```

*** =solution
```{r}
# Definisci la variabile vegas
vegas <- "Go!"
```

*** =sct
```{r}
test_object("vegas", incorrect_msg = "Assicurati di assegnare a `vegas` il corretto valore. Non dimenticarti che R è case sensitive!")
success_msg("Perfetto! Procedi all'esercizio successivo.")
```


--- type:NormalExercise xp:100 skills:1 key:aabd77b1a8
## Crea un vettore (2)

Concentriamoci sui vettori!

I vettori sono tra le strutture dati più utilizzate in R. All'inizio della tua esperienza in R farai sicuramente grande uso dei vettori.

Ma che cos'è un vettore? I vettori sono degli array monodimensionali che possono contenere dati di tipo numerico, character e logico. Puoi pensare ad un vettore come un contenitore di dati con più slots. Ad esempio, potresti immagazzinare in un vettore le vincite e perdite giornaliere al casinò.

In R puoi creare un vettore con la funzione combina [`c()`](http://www.rdocumentation.org/packages/base/functions/c). Puoi separare gli elementi contenuti nel vettore utilizzando una virgola. Per esempio:

```
numeric_vettore <- c(1, 2, 3)
character_vettore <- c("a", "b", "c")
```

Una volta dichiarati i vettori puoi utilizzarli per effettuare dei conti.

*** =instructions 
Completa il codice in modo che il vettore `boolean_vettore` contenga i seguenti tre elementi: `TRUE`, `FALSE` e `TRUE` (in questo specifico ordine). 

*** =hint 
Assegna `c(TRUE, FALSE, TRUE)` alla variabile `boolean_vettore` utilizzando l'operatore `<-`.

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
numeric_vettore <- c(1, 10, 49)
character_vettore <- c("a", "b", "c")

# Completa il codice per boolean_vettore
boolean_vettore <-
```

*** =solution
```{r}
numeric_vettore <- c(1, 10, 49)
character_vettore <- c("a", "b", "c")

# Completa il codice per boolean_vettore
boolean_vettore <- c(TRUE, FALSE, TRUE)
```

*** =sct
```{r}
msg <- "Non modificare il codice che dichiara i vettori `numeric_vettore` e `character_vettore`!"
test_object("numeric_vettore", undefined_msg = msg, incorrect_msg = msg)
test_object("character_vettore", undefined_msg = msg, incorrect_msg = msg)
test_object("boolean_vettore",
            incorrect_msg = "Assicurati di effettuare l'assegnazione di `boolean_vettore` correttamente. Utilizza `c(TRUE, FALSE, TRUE)`. Non utilizzare le virgolette per inserire i valori `TRUE` and `FALSE`! Infine, assicurati di mantenere lo stesso ordine richiesto dalle istruzioni.")
success_msg("Perfetto! Nota che aggiungere uno spazio dopo ogni virgola nella funzione `c()` migliora la leggibilità del codice. Facciamo un poco di pratica con la dichiarazione di alcuni vettori nel prossimo esercizio!")
```


--- type:NormalExercise xp:100 skills:1 key:2b0301eb5d
## Crea un vettore (3)

Dopo una settimana a Las Vegas e ancora nessuna Ferrari nel tuo garage, decidi che è tempo di utilizzare i tuoi superpoteri di data scientist.

Prima di effettuare una semplice analisi, decidi di raccogliere tutte le vincite e perdite della scorsa settimana: 

Per il vettore `vettore_poker`: 

- Lunedì hai vinto \$140
- Martedì hai perso \$50
- Mercoledì hai vinto \$20 
- Giovedì hai perso \$120
- Venerdì hai vinto \$240

Per il vettore `vettore_roulette`: 

- Lunedì hai vinto \$24
- Martedì hai perso \$50
- Mercoledì hai vinto \$100
- Giovedì hai perso \$350
- Venerdì hai vinto \$10

Hai giocato solo a poker e alla roulette siccome il tavolo di craps era sempre occupato. Per utilizzare questi dati in R, decidi di creare le variabili `vettore_poker` e `vettore_roulette`.

*** =instructions
Assegna le vittorie/perdite della roulette alla variabile `vettore_roulette`.

*** =hint
Per aiutarti in questo step, l'editor contiene già il codice per creare il vettore `vettore_poker`. Assegna il corretto valore a `vettore_roulette` sulla base dei numeri indicati nella consegna. Non dimenticare che i numeri riguardanti le perdite sono negativi!

*** =sample_code
```{r}
# Vincite a poker da lunedì a venerdì
vettore_poker <- c(140, -50, 20, -120, 240)

# Vincite alla roulette da lunedì a venerdì
vettore_roulette <-  
```

*** =solution
```{r}
# Vincite a poker da lunedì a venerdì
vettore_poker <- c(140, -50, 20, -120, 240)

# Vincite alla roulette da lunedì a venerdì
vettore_roulette <- c(-24, -50, 100, -350, 10)
```

*** =sct
```{r}
test_object("vettore_poker",
            incorrect_msg = "Assicurati di aver assegnato i valori corretti al vettore `vettore_poker`.")
test_object("vettore_roulette",
            incorrect_msg = "Assicurati di aver assegnato i valori corretti al vettore `vettore_roulette`. Assicurati inoltre di rispettare l'ordine dei valori!")
success_msg("Molto bene! Per visualizzare il contenuto dei vettori, ricordati che puoi sempre immettere semplicemente la variabile nella console e premere invio. Così facendo stamperai a video il vettore. Avanti al prossimo esercizio!")
```


--- type:NormalExercise xp:100 skills:1 key:29a7f48590
## Assegnare un nome ad un vettore

Dal punto di vista di un data analyst, è importante avere chiaro in mente che dati stai utilizzando. In particolare, è essenziale capire a cosa si riferisce ogni elemento.

Nell'esempio precedente abbiamo creato un vettore con le vincite della settimana. Ogni elemento del vettore si riferisce ad un giorno della settimana, però è difficile capire a quale giorno si riferisce un particolare elemento. Nel nostro esempio è relativamente semplice ricavare questa informazione ma nel caso in cui stessimo utilizzando, ad esempio, i giorni dell'anno potrebbe essere molto più complicato. Sarebbe quindi opportuno che il vettore contenesse anche questa informazione.

Puoi assegnare dei nomi agli elementi di un vettore utilizzando la funzione `names()`. Osserva questo esempio:

```
un_vettore <- c("Mario Bianchi", "Giocatore di Poker")
names(un_vettore) <- c("Nome", "Professione")
```

La prima riga di codice crea il vettore `un_vettore` mentre la seconda assegna i nomi agli elementi del vettore. Al primo elemento è assegnato il nome `Nome`, mentre il secondo elemento è etichettato come `Professione`. Stampando il contenuto nella console si ottiene il seguente output:

```
          Nome     Professione 
"Mario Bianchi"   "Giocatore di Poker" 
```

*** =instructions
- Il codice sulla destra assegna un nome agli elementi del vettore `vettore_poker` utilizzando i giorni della settimana. Aggiungi il codice per effettuare la stessa operazione sul vettore `vettore_roulette`.

*** =hint
Puoi utilizzare `names(vettore_roulette)` per assegnare i nomi del vettore `vettore_roulette`. Assicurati di utilizzare lo stesso vettore con i giorni della settimana come nomi. Ancora una volta, ricorda che R è case sensitive!

*** =sample_code
```{r}
# Vincite a poker da lunedì a venerdì
vettore_poker <- c(140, -50, 20, -120, 240)

# Vincite alla roulette da lunedì a venerdì
vettore_roulette <- c(-24, -50, 100, -350, 10)

# Assegna i giorni della settimana come nomi a vettore_poker
names(vettore_poker) <- c("Lunedì", "Martedì", "Mercoledì", "Giovedì", "Venerdì")

# Assegna i giorni della settimana come nomi a vettore_roulette
```

*** =solution
```{r}
# Vincite a poker da lunedì a venerdì
vettore_poker <- c(140, -50, 20, -120, 240)

# Vincite alla roulette da lunedì a venerdì
vettore_roulette <- c(-24, -50, 100, -350, 10)

# Assegna i giorni della settimana come nomi a vettore_poker
names(vettore_poker) <- c("Lunedì", "Martedì", "Mercoledì", "Giovedì", "Venerdì")

# Assegna i giorni della settimana come nomi a vettore_roulette
names(vettore_roulette) <- c("Lunedì", "Martedì", "Mercoledì", "Giovedì", "Venerdì")
```

*** =sct
```{r}
test_object("vettore_poker",
            incorrect_msg = "Non cambiare i valori dentro `vettore_poker`; Sono già stati inseriti per te.")
test_object("vettore_roulette",
            incorrect_msg = "Non cambiare i valori dentro `vettore_roulette`; Sono già stati inseriti per te.")
test_object("vettore_poker",
            eq_condition = "equal",
            incorrect_msg = "Non cambiare il codice che assegna i nomi a `vettore_poker`; Concentrati su `vettore_roulette`!")
test_object("vettore_roulette",
            eq_condition = "equal",
            incorrect_msg = "Assicurati che l'assegnazione dei nomi al vettore `vettore_roulette` sia stata effettuata in modo corretto. Utilizza lo stesso vettore utilizzato per assegnare i nomi a `vettore_poker`.")
success_msg("Ben fatto! Prosegui al prossimo esercizio.")
```


--- type:NormalExercise xp:100 skills:1 key:33b17b05db
## Assegnare un nome ad un vettore (2)

Se vuoi diventare un buon programmatore devi diventare pigro. (Se sei già pigro, è molto probabile che tu sia già uno di quei talenti nell'ambito della programmazione in attesa di esprimerti).

Negli esercizi precedenti, probabilmente, hai già provato un pò di sana noia e frustrazione nel digitare più e più volte le stesse informazioni come ad esempio i giorni della settimana. Tuttavia, quando osservi il problema da un'altra prospettiva, esiste un modo più efficiente di effettuare queste operazioni, per esempio assegnando il vettore dei giorni della settimana ad una **variabile**! 

Come hai fatto con i dati riguardanti il poker e la roulette, puoi creare una variabile che contenga i giorni della settimana. In questo modo puoi utilizzarli più volte senza dover digitare nuovamente tutto il contenuto della variabile.

*** =instructions
- Una variabile `vettore_giorni` che contiene i giorni della settimana è già stata creata per te.
- Usa `vettore_giorni` per assegnare i nomi ai vettori `vettore_poker` e `vettore_roulette`.

*** =hint
Utilizza `names(vettore_poker) <- vettore_giorni` per assegnare i nomi agli elementi di `vettore_poker`. Fai la stessa cosa per `vettore_roulette`.

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# Vincite a poker da lunedì a venerdì
vettore_poker <- c(140, -50, 20, -120, 240)

# Vincite alla roulette da lunedì a venerdì
vettore_roulette <- c(-24, -50, 100, -350, 10)

# La variabile vettore_giorni
vettore_giorni <- c("Lunedì", "Martedì", "Mercoledì", "Giovedì", "Venerdì")
 
# Assegna il nome dei giorni a vettore_roulette e vettore_poker
names(vettore_poker) <-   
names(vettore_roulette) <-
```

*** =solution
```{r}
# Vincite a poker da lunedì a venerdì
vettore_poker <- c(140, -50, 20, -120, 240)

# Vincite alla roulette da lunedì a venerdì
vettore_roulette <- c(-24, -50, 100, -350, 10)

# La variabile vettore_giorni
vettore_giorni <- c("Lunedì", "Martedì", "Mercoledì", "Giovedì", "Venerdì")

# Assegna il nome dei giorni a vettore_roulette e vettore_poker
names(vettore_poker) <- vettore_giorni
names(vettore_roulette) <- vettore_giorni
```

*** =sct
```{r}
msg <- "Non cambiare le variabili predefinite `vettore_poker`, `vettore_roulette` e `vettore_giorni`."
test_object("vettore_poker", undefined_msg = msg, incorrect_msg = msg)
test_object("vettore_roulette", undefined_msg = msg, incorrect_msg = msg)
test_object("vettore_giorni", undefined_msg = msg, incorrect_msg = msg)

test_object("vettore_poker",
            eq_condition = "equal",
            incorrect_msg = "Assicurati di assegnare `vettore_giorni` ai nomi di `vettore_poker`.")
test_object("vettore_roulette",
            eq_condition = "equal",
            incorrect_msg = "Assicurati di assegnare `vettore_giorni` ai nomi di `vettore_roulette`.")
success_msg("Ottimo! Un consiglio: cerca di evitare di duplicare il codice ogni volta che puoi. Avanti alla prossima sezione dove imparerai ad effettuare operazioni aritmentiche con i vettori!")
```


--- type:NormalExercise xp:100 skills:1 key:6c16c77630
## Calcola le vincite totali

Ora che i dati sulle vincite sono stati immagazzinati in vettori puoi iniziare a fare alcuni conti. 

Diciamo che sei interessato a ottenere le seguenti informazioni:

- Quanto è stato il profitto/ la perdita medio/a per giorno della settimana?

- Hai perso soldi in totale durante la settimana?

- Stai vincendo/perdendo alla roulette/poker?

Per ottenere le risposte, devi effettuare alcune operazioni aritmentiche sui vettori.

Nota che in R se effettui la somma tra due vettori, la somma viene effettuata elemento per elemento. Per esempio, le tre linee di codice seguenti sono completamente equivalenti:

```
c(1, 2, 3) + c(4, 5, 6)
c(1 + 4, 2 + 5, 3 + 6)
c(5, 7, 9)
```

Puoi anche effettuare la somma con delle variabili alle quali hai precedentemente assegnato i vettori

```
a <- c(1, 2, 3) 
b <- c(4, 5, 6)
c <- a + b
```

*** =instructions
- Somma le variabili `vettore_A` e `vettore_B` e assegna il risultato a `vettore_totale`.
- Controlla il risultato stampando a video `vettore_totale`.

*** =hint
Usa l'operatore `+` per sommare `vettore_A` e `vettore_B`. Usa `<-` per assegnare il risultato a `vettore_totale`.

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
vettore_A <- c(1, 2, 3)
vettore_B <- c(4, 5, 6)

# Somma vettore_A e vettore_B
vettore_totale <- 
  
# Stampa a video vettore_totale

```

*** =solution
```{r}
vettore_A <- c(1, 2, 3)
vettore_B <- c(4, 5, 6)

# Somma vettore_A e vettore_B
vettore_totale <- vettore_A + vettore_B

# Stampa a video vettore_totale
vettore_totale
```

*** =sct
```{r}
msg <- "Non cambiare il contenuto di `vettore_A` e `vettore_B`!"
test_object("vettore_A", undefined_msg = msg, incorrect_msg = msg)
test_object("vettore_B", undefined_msg = msg, incorrect_msg = msg)
test_object("vettore_totale", incorrect_msg = "Assicurati che `vettore_totale` contenga la somma di `vettore_A` e `vettore_B`.")
test_output_contains("vettore_totale", incorrect_msg = "Non dimenticare di stampare a video `vettore_totale`! Scrivi semplicemente `vettore_totale` su una nuova riga.")
success_msg("Ottimo lavoro! Prosegui al prossimo esercizio.")
```


--- type:NormalExercise xp:100 skills:1 key:49ee1d095a
## Calcola le vincite totali (2)

Ora sai come R effettua operazioni aritmetiche con i vettori. Adesso calcola a quanto ammonta il profitto o la perdita totale per giorno della settimana. Il profitto totale giornaliero è la somma dei profitti e delle perdite che hai realizzato a poker sommato al risultato giornaliero del gioco della roulette.

In R, questa operazione corrisponde semplicemente alla somma di `vettore_roulette` e `vettore_poker`.

*** =instructions
Assegna alla variabile `totale_giornaliero` quanto hai vinto o perso ogni giorno in totale (combinando poker e roulette).

*** =hint
In modo del tutto simile all'esercizio precedente, assegna la somma dei due vettori ad una nuova variabile chiamata `totale_giornaliero`.

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# Vincite al poker e alla roulette da lunedì a venerdì
vettore_poker <- c(140, -50, 20, -120, 240)
vettore_roulette <- c(-24, -50, 100, -350, 10)
vettore_giorni <- c("Lunedì", "Martedì", "Mercoledì", "Giovedì", "Venerdì")
names(vettore_poker) <- vettore_giorni
names(vettore_roulette) <- vettore_giorni

# Assegna a totale_giornaliero quanto hai vinto/perso ogni giorno
totale_giornaliero <- 
```

*** =solution
```{r}
# Vincite al poker e alla roulette da lunedì a venerdì
vettore_poker <- c(140, -50, 20, -120, 240)
vettore_roulette <- c(-24, -50, 100, -350, 10)
vettore_giorni <- c("Lunedì", "Martedì", "Mercoledì", "Giovedì", "Venerdì")
names(vettore_poker) <- vettore_giorni
names(vettore_roulette) <- vettore_giorni

# Assegna a totale_giornaliero quanto hai vinto/perso ogni giorno
totale_giornaliero <- vettore_poker + vettore_roulette
```

*** =sct
```{r}
msg = "Non cambiare nulla nella dichiarazione e assegnamento nomi di `vettore_poker` e `vettore_roulette`."
test_object("vettore_giorni", undefined_msg = msg, incorrect_msg = msg)
test_object("vettore_poker", eq_condition = "equal", undefined_msg = msg, incorrect_msg = msg)
test_object("vettore_roulette", eq_condition = "equal", undefined_msg = msg, incorrect_msg = msg)

test_object("totale_giornaliero", incorrect_msg = "Assicurati di assegnare la somma di `vettore_poker` e `vettore_roulette` a `totale_giornaliero`.")

success_msg("Eccellente! Continua all'esercizio successivo.")
```


--- type:NormalExercise xp:100 skills:1 key:113f861114
## Calcola le vincite totali (3)

Sulla base delle analisi effettuate precedentemente, sembra che tu abbia ottenuto un risultato positivo in alcuni giorni e negativo in altri. Questo non è chiaramente il risultato ottimale atteso, e ti chiedi se globalmente ci sia stato un profitto o una perdita.

Per rispondere a questa domanda, puoi utilizzare la funzione [`sum()`](http://www.rdocumentation.org/packages/base/functions/sum). Questa funzione accetta un vettore come parametro e calcola la somma di tutti gli elementi di quel vettore. Per esempio, per calcolare il totale vinto/perso a poker, puoi fare nel modo seguente:

```
totale_poker <- sum(vettore_poker)
```

*** =instructions
- Calcola il totale che hai vinto/perso al gioco della roulette e assegna il risultato alla variabile `totale_roulette`.
- Ora che hai il risultato totale per entrambi i giochi, puoi calcolare facilmente il risultato globale della settimana `totale_settimana` (che è la somma di tutte le perdite e profitti della settimana).
- Stampa il contenuto di `totale_settimana` a video.

*** =hint
Usa la funzione [`sum()`](http://www.rdocumentation.org/packages/base/functions/sum) per sommare tutti gli elementi del vettore `vettore_roulette`. `totale_settimana` è la somma complessiva di `vettore_roulette` e `vettore_poker`.

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# Vincite a poker e alla roulette da lunedì a venerdì
vettore_poker <- c(140, -50, 20, -120, 240)
vettore_roulette <- c(-24, -50, 100, -350, 10)
vettore_giorni <- c("Lunedì", "Martedì", "Mercoledì", "Giovedì", "Venerdì")
names(vettore_poker) <- vettore_giorni
names(vettore_roulette) <- vettore_giorni

# Vincite complessive a poker
totale_poker <- sum(vettore_poker)

# Vincite totali alla roulette
totale_roulette <-  

# Vincite complessive
totale_settimana <- 

# Stampa total_week
  
```

*** =solution
```{r}
# Vincite a poker e alla roulette da lunedì a venerdì
vettore_poker <- c(140, -50, 20, -120, 240)
vettore_roulette <- c(-24, -50, 100, -350, 10)
vettore_giorni <- c("Lunedì", "Martedì", "Mercoledì", "Giovedì", "Venerdì")
names(vettore_poker) <- vettore_giorni
names(vettore_roulette) <- vettore_giorni

# Vincite complessive a poker
totale_poker <- sum(vettore_poker)

# Vincite totali alla roulette
totale_roulette <-  sum(vettore_roulette)

# Vincite complessive
totale_settimana <- total_roulette + total_poker

# Stampa total_week
totale_settimana
```

*** =sct
```{r}
msg = "Non cambiare nulla nella dichiarazione e nell'assegnazione dei nomi di `vettore_poker` e `vettore_roulette`."
test_object("vettore_giorni", undefined_msg = msg, incorrect_msg = msg)
test_object("vettore_poker", eq_condition = "equal", undefined_msg = msg, incorrect_msg = msg)
test_object("vettore_roulette", eq_condition = "equal", undefined_msg = msg, incorrect_msg = msg)
test_object("totale_poker", 
            incorrect_msg = "Assicurati di assegnare a `totale_poker` la somma del vettore `vettore_poker`.")
test_object("totale_roulette",
            incorrect_msg = "Assicurati di assegnare a `totale_roulette` la somma del vettore `vettore_roulette`.")
test_object("totale_settimana",
            incorrect_msg = "Assicurati di assegnare a `totale_settimana` le somme dei due vettori: `totale_roulette` e `totale_poker`.")

test_output_contains("totale_settimana", incorrect_msg = "Non dimenticarti di digitare `total_week` su una nuova riga per stampare il contenuto della variabile.")
success_msg("Ottimo. Questa però è una cattiva notizia...")
```


--- type:NormalExercise xp:100 skills:1 key:4be7ee1045
## Confrontare le vincite totali

Oops, sembra che tu stia perdendo invece di vincere. Tempo di rivedere la tua strategia! Questo però richiederà analisi più dettagliate...

Dopo una breve sessione di brainstorming nella jacuzzi dell'hotel, giungi alla conclusione che è possibile che le tue abilità nel gioco della roulette non siano al livello di quelle del gioco del poker. Può darsi quindi che le tue vincite totali a poker siano maggiori (o `>` ) rispetto a quelle della roulette.

*** =instructions
- Calcola `totale_poker` e `totale_roulette` come nell'esercizio precedente. Utilizza la funzione `sum()` due volte.
- Controlla se le vincite totali al poker sono maggiori di quelle alla roulette utilizzando il simbolo di disuguaglianza opportuno. Stampa il risultato del confronto. Cosa puoi concludere? Dovresti focalizzarti più sul poker o sulla roulette?

*** =hint
- In parte hai già calcolato la risposta a questo esercizio nell'esercizio precedente!
- Per verificare se 6 è maggiore di 5, digita `6 > 5`. Questa operazione ritornerà un valore logico (`TRUE` oppure `FALSE`).

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# Vincite al poker e alla rulette da lunedì a venerdì:
vettore_poker <- c(140, -50, 20, -120, 240)
vettore_roulette <- c(-24, -50, 100, -350, 10)
vettore_giorni <- c("Lunedì", "Martedì", "Mercoledì", "Giovedì", "Venerdì")
names(vettore_poker) <- vettore_giorni
names(vettore_roulette) <- vettore_giorni

# Calcola vincite totali al poker e alla roulette
totale_poker <-
totale_roulette <-

# Controlla se hai relizzato vincite superiori al poker rispetto alla roulette

```

*** =solution
```{r}
# Vincite al poker e alla rulette da lunedì a venerdì:
vettore_poker <- c(140, -50, 20, -120, 240)
vettore_roulette <- c(-24, -50, 100, -350, 10)
vettore_giorni <- c("Lunedì", "Martedì", "Mercoledì", "Giovedì", "Venerdì")
names(vettore_poker) <- vettore_giorni
names(vettore_roulette) <- vettore_giorni

# Calcola vincite totali al poker e alla roulette
totale_poker <- sum(vettore_poker)
totale_roulette <- sum(vettore_roulette)

# Controlla se hai relizzato vincite superiori al poker rispetto alla roulette
totale_poker > totale_roulette
```

*** =sct
```{r}
msg <- "Non cambiare nulla nella dichiarazione e assegnazione nomi dei vettori `vettore_poker` e `vettore_roulette`."
test_object("vettore_giorni", undefined_msg = msg, incorrect_msg = msg)
test_object("vettore_poker", eq_condition = "equal", undefined_msg = msg, incorrect_msg = msg)
test_object("vettore_roulette", eq_condition = "equal", undefined_msg = msg, incorrect_msg = msg)

test_object("totale_poker", 
            incorrect_msg = "Assicurati di assegnare a `total_poker` la somma di `vettore_poker`. Usa `sum()`.")
test_object("totale_roulette",
            incorrect_msg = "Assicurati di assegnare a `total_roulette` la somma di `vettore_roulette`. Usa `sum()`.")
test_output_contains("totale_poker > totale_roulette",
                     incorrect_msg = "Hai effettuato il confronto in modo corretto? Per verificare se `total_poker` è maggiore di `total_roulette`, puoi utilizzare il codice seguente `total_poker > total_roulette`.")
success_msg("Ottimo lavoro! Procedi al prossimo esercizio.")
```


--- type:NormalExercise xp:100 skills:1 key:bb0eb55d93
## Selezione elementi dei vettori

La tua idea sembra essere corretta. Sembra che il poker ti si addica di più della roulette. 

Un'altra possibilità da investigare è che la tua performance sia migliore all'inizio della settimana e peggiore verso il weekend. Effettivamente verso il fine settimana hai bevuto qualche cocktail che potrebbe aver peggiorato la situazione... 

Per analizzare questa possibilità, devi focalizzarti solo su una parte del vettore `vettore_totale`. In altre parole, il tuo obiettivo è di selezionare elementi specifici del vettore. Per selezionare solo alcuni elementi di un vettore (e successivamente di matrici, data frames, ...) devi utilizzare le parentesi quadre. Tra le parentesi quadre devi indicare l'indice dell'elemento da selezionare. Per esempio, per selezionare il primo elemento, basta digitare `vettore_poker[1]`, mentre per selezionare il secondo elemento occorre utilizzare `vettore_poker[2]`, etc... Nota che il primo elemento ha indice 1 e non 0 come accade in molti altri linguaggi di programmazione.

*** =instructions
Assegna i risultati delle giocate effettuate il Mercoledì alla variabile `poker_mercoledi`.

*** =hint
I dati riguardanti mercoledì sono il terzo elemento di `vettore_poker`, quindi possono essere ottenuti digitando `vettore_poker[3]`.

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# Vincite al poker e alla rulette da lunedì a venerdì:
vettore_poker <- c(140, -50, 20, -120, 240)
vettore_roulette <- c(-24, -50, 100, -350, 10)
vettore_giorni <- c("Lunedì", "Martedì", "Mercoledì", "Giovedì", "Venerdì")
names(vettore_poker) <- vettore_giorni
names(vettore_roulette) <- vettore_giorni

# Definisci una nuova variabile sulla base della selezione
poker_mercoledi <- 
```

*** =solution
```{r}
# Vincite al poker e alla rulette da lunedì a venerdì:
vettore_poker <- c(140, -50, 20, -120, 240)
vettore_roulette <- c(-24, -50, 100, -350, 10)
vettore_giorni <- c("Lunedì", "Martedì", "Mercoledì", "Giovedì", "Venerdì")
names(vettore_poker) <- vettore_giorni
names(vettore_roulette) <- vettore_giorni

# Definisci una nuova variabile sulla base della selezione
poker_mercoledi <- vettore_poker[3]
```

*** =sct
```{r}
msg = "Non cambiare nulla nella dichiarazione e assegnazione nomi dei vettori `vettore_poker` e `vettore_roulette`."
test_object("vettore_giorni", undefined_msg = msg, incorrect_msg = msg)
test_object("vettore_poker", eq_condition = "equal", undefined_msg = msg, incorrect_msg = msg)
test_object("vettore_roulette", eq_condition = "equal", undefined_msg = msg, incorrect_msg = msg)
test_object("poker_mercoledi", 
            undefined_msg = "Assicurati di definire una variabile `poker_mercoledi`.",
            incorrect_msg = "Sembra che `poker_mercoledi` il corretto valore selezionato da `vettore_poker`.")
success_msg("Eccellente! R permette anche di selezionari più valori. Impara come nell'esercizio successivo!")
```


--- type:NormalExercise xp:100 skills:1 key:ce93f7a2ff
## Selezione elementi dei vettori (2)

Perchè non analizzare i risultati infrasettimanali? 

Per selezionari più elementi da un vettore inserisci nelle parentesi quadre gli indici degli elementi che intendi selezionare. Puoi unire gli indici utilizzando la funzione `c()` già vista negli esercizi precedenti. Per esempio: supponi di voler selezionare il primo e il quinto giorno della settimana, per fare questo puoi utilizzare il vettore `c(1, 5)` tra le parentesi quadre. Il codice riportato di seguito seleziona il primo e il quinto elemento dal vettore `vettore_poker`:

```
vettore_poker[c(1, 5)]
```

*** =instructions
Assegna i risultati di Martedì, Mercoledì e Giovedì alla variabile `poker_midweek`.

*** =hint
Usa il vettore `c(2, 3, 4)` tra parentesi quadre per selezionare gli elementi corretti di `vettore_poker`.

*** =pre_exercise_code
```{r}
# no pec
``` 

*** =sample_code
```{r}
# Vincite al poker e alla rulette da lunedì a venerdì:
vettore_poker <- c(140, -50, 20, -120, 240)
vettore_roulette <- c(-24, -50, 100, -350, 10)
vettore_giorni <- c("Lunedì", "Martedì", "Mercoledì", "Giovedì", "Venerdì")
names(vettore_poker) <- vettore_giorni
names(vettore_roulette) <- vettore_giorni

# Definisci una nuova variabile sulla base della selezione
poker_midweek <- 
```

*** =solution
```{r}
# Vincite al poker e alla rulette da lunedì a venerdì:
vettore_poker <- c(140, -50, 20, -120, 240)
vettore_roulette <- c(-24, -50, 100, -350, 10)
vettore_giorni <- c("Lunedì", "Martedì", "Mercoledì", "Giovedì", "Venerdì")
names(vettore_poker) <- vettore_giorni
names(vettore_roulette) <- vettore_giorni

# Definisci una nuova variabile sulla base della selezione
poker_midweek <- vettore_poker[c(2, 3, 4)]
```

*** =sct
```{r}
msg = "Non cambiare nulla nella dichiarazione e assegnazione nomi dei vettori `vettore_poker` e `vettore_roulette`."
test_object("vettore_giorni", undefined_msg = msg, incorrect_msg = msg)
test_object("vettore_poker", eq_condition = "equal", undefined_msg = msg, incorrect_msg = msg)
test_object("vettore_roulette", eq_condition = "equal", undefined_msg = msg, incorrect_msg = msg)
test_object("poker_midweek", 
            incorrect_msg = "Sembra che `poker_midweek` non contenga i valori corretti selezionati da `vettore_poker`. Puoi utilizzare il vettore `c(2, 3, 4)` dentro le parentesi quadre per ottenere il risultato desiderato.")
success_msg("Molto bene! Continua con l'esercizio successivo per specializzarti nella selezione di elementi dai vettori!");
```


--- type:NormalExercise xp:100 skills:1 key:f3419b5483
## Selezione elementi dei vettori (3)

Selezionare più elementi da `vettore_poker` utilizzando `c(2, 3, 4)` non è molto comodo. Molti programmatori sono pigri per natura e così hanno creato un metodo più semplice e veloce: `c(2, 3, 4)` può essere abbreviato utilizzando `2:4`, ossia un comando che genera un vettore con tutti i numeri naturali compresi tra 2 e 4, estremi inclusi.

Quindi, un altro modo per trovare il risultato infrasettimanale è `vettore_poker[2:4]`. Nota come `2:4` sia posizionato dentro le parentesi quadre per selezionare gli elementi da 2 a 4 inclusi.

*** =instructions
Assegna a `vettore_selezione_roulette` i risultati ottenuti alla roulette da Martedì fino a Venerdì; utilizza `:` se ti facilita rispetto al metodo utilizzato precedentemente.

*** =hint
Assegna una selezione di elementi dal vettore `vettore_roulette` al vettore `vettore_selezione_roulette` semplicemente mettendo `2:5` tra la parentesi quadre.

*** =sample_code
```{r}
# Vincite al poker e alla rulette da lunedì a venerdì:
vettore_poker <- c(140, -50, 20, -120, 240)
vettore_roulette <- c(-24, -50, 100, -350, 10)
vettore_giorni <- c("Lunedì", "Martedì", "Mercoledì", "Giovedì", "Venerdì")
names(vettore_poker) <- vettore_giorni
names(vettore_roulette) <- vettore_giorni

# Definisci una nuova variabile sulla base della selezione
vettore_selezione_roulette <- 
```

*** =solution
```{r}
# Vincite al poker e alla rulette da lunedì a venerdì:
vettore_poker <- c(140, -50, 20, -120, 240)
vettore_roulette <- c(-24, -50, 100, -350, 10)
vettore_giorni <- c("Lunedì", "Martedì", "Mercoledì", "Giovedì", "Venerdì")
names(vettore_poker) <- vettore_giorni
names(vettore_roulette) <- vettore_giorni

# Definisci una nuova variabile sulla base della selezione
vettore_selezione_roulette <- vettore_roulette[2:5]
```

*** =sct

```{r}
msg = "Non cambiare nulla nella dichiarazione e assegnazione nomi dei vettori `vettore_poker` e `vettore_roulette`."
test_object("vettore_giorni", undefined_msg = msg, incorrect_msg = msg)
test_object("vettore_poker", eq_condition = "equal", undefined_msg = msg, incorrect_msg = msg)
test_object("vettore_roulette", eq_condition = "equal", undefined_msg = msg, incorrect_msg = msg)
test_object("vettore_selezione_roulette", 
            undefined_msg = "Assicurati di definire una variabile `vettore_selezione_roulette`.",
            incorrect_msg = "Sembra che `vettore_selezione_roulette` non contenga gli elementi corretti selezionati da `vettore_roulette`. Assicurati di utilizzare gli indici corretti.")
success_msg("Fantastico! L'operatore due punti è estremamente utile ed è spesso utilizzato nella programmazione in R, quindi ricordatelo bene! Procedi al prossimo esercizio.")
```


--- type:NormalExercise xp:100 skills:1 key:0ed272a349
## Selezione elementi dei vettori (4)

Un ulteriore modo per risolvere l'esercizio precedente consiste nell'utilizzare i nomi degli elementi del vettore (nel nostro caso: Lunedì, Martedì, ...) invece degli indici di posizione. Ad esempio, 

```
vettore_poker["Lunedì"]
```

selezionerà il primo elemento di `vettore_poker` siccome `"Lunedì"` è il nome del primo elemento.

Allo stesso modo, come hai fatto precedentemente con i vettori numerici, puoi utilizzare anche i nomi degli elementi per selezionare più elementi in una sola volta. Ad esempio:

```
vettore_poker[c("Lunedì","Martedì")]
```

*** =instructions
- Seleziona i primi tre elementi di `vettore_poker` utilizzando i loro nomi: `"Lunedì"`, `"Martedì"` and `"Mercoledì"`. Assegna il risultato alla variabile `poker_inizio_settimana`.
- Calcola la media della variabile `poker_inizio_settimana` utilizzando la funzione [`mean()`](http://www.rdocumentation.org/packages/base/functions/mean). Stampa il risultato a video per visualizzarlo.

*** =hint
- Puoi utilizzare `c("Lunedì", "Martedì", "Mercoledì")` dentro le parentesi quadre per ottenere gli elementi di `vettore_poker` richiesti.
- Puoi utilizzare `mean(poker_inizio_settimana)` per calcolare la media degli elementi in `poker_inizio_settimana`. Non è necessario effettuare la media di tutti i giorni, è sufficiente effettuare la media dei primi tre.

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# Vincite al poker e alla rulette da lunedì a venerdì:
vettore_poker <- c(140, -50, 20, -120, 240)
vettore_roulette <- c(-24, -50, 100, -350, 10)
vettore_giorni <- c("Lunedì", "Martedì", "Mercoledì", "Giovedì", "Venerdì")
names(vettore_poker) <- vettore_giorni
names(vettore_roulette) <- vettore_giorni

# Seleziona i risultati del poker di Lunedì, Martedì e Mercoledì
poker_inizio_settimana <- 
  
# Calcola la media degli elementi in poker_inizio_settimana

```

*** =solution
```{r}
# Vincite al poker e alla rulette da lunedì a venerdì:
vettore_poker <- c(140, -50, 20, -120, 240)
vettore_roulette <- c(-24, -50, 100, -350, 10)
vettore_giorni <- c("Lunedì", "Martedì", "Mercoledì", "Giovedì", "Venerdì")
names(vettore_poker) <- vettore_giorni
names(vettore_roulette) <- vettore_giorni

# Seleziona i risultati del poker di Lunedì, Martedì e Mercoledì
poker_inizio_settimana <- vettore_poker[c("Lunedì", "Martedì", "Mercoledì")]
  
# Calcola la media degli elementi in poker_inizio_settimana
mean(poker_inizio_settimana)
```

*** =sct
```{r}
msg = "Non cambiare nulla nella dichiarazione e assegnazione nomi dei vettori `vettore_poker` e `vettore_roulette`."
test_object("vettore_giorni", undefined_msg = msg, incorrect_msg = msg)
test_object("vettore_poker", eq_condition = "equal", undefined_msg = msg, incorrect_msg = msg)
test_object("vettore_roulette", eq_condition = "equal", undefined_msg = msg, incorrect_msg = msg)
test_object("poker_inizio_settimana", 
            incorrect_msg = "Sembra che `poker_inizio_settimana` non contenga i primi tre elementi di `vettore_poker`. Puoi utilizzare `c(\"Lunedì\", \"Martedì\", \"Mercoledì\")` dentro le parentesi quadre per ottenere il risultato desiderato.")
test_output_contains("mean(poker_inizio_settimana)", incorrect_msg = "Hai calcolato la media dei valori contenuti in `poker_inizio_settimana` e li hai stampati? Utilizza `mean(poker_inizio_settimana)`.")
success_msg("Ottimo lavoro! Oltre a partizionare i vettori utilizzando gli indici e i nomi, puoi effettuare la stessa operazione effettuando dei confronti. L'esercizio successivo ti mostrerà come fare!");
```


--- type:NormalExercise xp:100 skills:1 key:9c28970fd6
## Selezione attraverso il confronto - Step 1

Utilizzando gli operatori di confronto puoi approcciare il problema affrontato precedentemente in modo differente.

Gli operatori di confronto (logici) in R sono i seguenti:

- `<` per indicare minore di
- `>` per indicare maggiore di
- `<=` per indicare minore o uguale a
- `>=` per indicare maggiore o uguale a
- `==` per indicare l'uguaglianza
- `!=` per indicare la disuguaglianza

Come visto nel capitolo precedente, la riga di codice `6 > 5` ritorna `TRUE` (vero). R permette di utilizzare gli operatori di confronto anche sui vettori. Ad esempio:

```
> c(4, 5, 6) > 5
[1] FALSE FALSE TRUE
```

Questo comando verifica, per ogni elemento del vettore, se la condizione affermata dall'operatore di confronto è `TRUE` oppure `FALSE`.

*** =instructions
- Controlla quali elementi di `vettore_poker` sono positivi (i.e. > 0) e assegna il risultato alla variabile `vettore_selezione`. 
- Stampa `vettore_selezione` in modo da poterlo visualizzare. Il risultato stampato ti permette di individuare se hai vinto (`TRUE`) o perso (`FALSE`) in ciascun giorno.

*** =hint
Per verificare in quali giorni i risultati ottenuti a poker sono positivi, R deve verificare la disuguaglianza su ogni elemento del vettore `vettore_poker`. Per effettuare il confronto utilizza una espressione simile a questa: `un_vettore > 0`.

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# Vincite al poker e alla rulette da lunedì a venerdì:
vettore_poker <- c(140, -50, 20, -120, 240)
vettore_roulette <- c(-24, -50, 100, -350, 10)
vettore_giorni <- c("Lunedì", "Martedì", "Mercoledì", "Giovedì", "Venerdì")
names(vettore_poker) <- vettore_giorni
names(vettore_roulette) <- vettore_giorni

# In quali giorni hai vinto a poker?
vettore_selezione <- 
  
# Stampa vettore_selezione

```

*** =solution
```{r}
# Vincite al poker e alla rulette da lunedì a venerdì:
vettore_poker <- c(140, -50, 20, -120, 240)
vettore_roulette <- c(-24, -50, 100, -350, 10)
vettore_giorni <- c("Lunedì", "Martedì", "Mercoledì", "Giovedì", "Venerdì")
names(vettore_poker) <- vettore_giorni
names(vettore_roulette) <- vettore_giorni

# In quali giorni hai vinto a poker?
vettore_selezione <- vettore_poker > 0
  
# Stampa vettore_selezione
vettore_selezione
```

*** =sct
```{r}
msg <- "Non cambiare nulla nella dichiarazione e assegnazione nomi dei vettori `vettore_poker` e `vettore_roulette`."
test_object("vettore_giorni", undefined_msg = msg, incorrect_msg = msg)
test_object("vettore_poker", eq_condition = "equal", undefined_msg = msg, incorrect_msg = msg)
test_object("vettore_roulette", eq_condition = "equal", undefined_msg = msg, incorrect_msg = msg)

test_object("vettore_selezione", incorrect_msg = "Sembra che `vettore_selezione` non contenga il risultato corretto. Ricorda che R effettua le operazioni sui vettori elemento per elemento.")
test_output_contains("vettore_selezione", incorrect_msg = "Non dimenticarti di stampare `vettore_selezione` scrivendo il nome della variabile su una riga nuova.")
success_msg("Ottimo!")
```


--- type:NormalExercise xp:100 skills:1 key:682d16f025
## Selezione attraverso il confronto - Step 2

L'utilizzo degli operatori di confronto renderà la tua vita di data scientist molto più semplice. Invece di selezionare un sottoinsieme di giorni, puoi chiedere a R di selezionare solo quei giorni dove hai relizzato una vincita al poker.

Nell'esercizio precedente hai utilizzato `vettore_selezione <- vettore_poker > 0` per trovare i giorni vincenti. Adesso vorresti sapere non solo i giorni nei quali hai vinto ma anche l'ammontare della vincita.

Puoi selezionare gli elementi desiderati inserendo semplicemente il `vettore_selezione` tra le parentesi quadre che seguono il vettore `vettore_poker`:

```
vettore_poker[vettore_selezione]
```

R sa già cosa fare quando inserisci un vettore logico nelle parentesi quadre: selezionerà solo gli elementi che corrispondono a `TRUE` nel vettore `vettore_selezione`.

*** =instructions
Utilizza `vettore_selezione` nelle parentesi quadre per assegnare l'ammontare delle vincite dei giorni vincenti alla variabile `poker_giorni_vincenti`.

*** =hint
Utilizza `vettore_poker[vettore_selezione]` per selezionare gli elementi desiderati da `vettore_poker`, e assegna il risultato a `poker_giorni_vincenti`.

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# Vincite al poker e alla rulette da lunedì a venerdì:
vettore_poker <- c(140, -50, 20, -120, 240)
vettore_roulette <- c(-24, -50, 100, -350, 10)
vettore_giorni <- c("Lunedì", "Martedì", "Mercoledì", "Giovedì", "Venerdì")
names(vettore_poker) <- vettore_giorni
names(vettore_roulette) <- vettore_giorni

# In quali giorni hai vinto a poker?
vettore_selezione <- vettore_poker > 0

# Seleziona questi giorni dal vettore_poker
poker_giorni_vincenti <- 
```

*** =solution
```{r}
# Vincite al poker e alla rulette da lunedì a venerdì:
vettore_poker <- c(140, -50, 20, -120, 240)
vettore_roulette <- c(-24, -50, 100, -350, 10)
vettore_giorni <- c("Lunedì", "Martedì", "Mercoledì", "Giovedì", "Venerdì")
names(vettore_poker) <- vettore_giorni
names(vettore_roulette) <- vettore_giorni

# In quali giorni hai vinto a poker?
vettore_selezione <- vettore_poker > 0

# Seleziona questi giorni dal vettore_poker
poker_giorni_vincenti <- vettore_poker[vettore_selezione]
```

*** =sct
```{r}
msg = "Non cambiare nulla nella dichiarazione e assegnazione nomi dei vettori `vettore_poker` e `vettore_roulette`."
test_object("vettore_giorni", undefined_msg = msg, incorrect_msg = msg)
test_object("vettore_poker", eq_condition = "equal", undefined_msg = msg, incorrect_msg = msg)
test_object("vettore_roulette", eq_condition = "equal", undefined_msg = msg, incorrect_msg = msg)
test_object("vettore_selezione", incorrect_msg = "Non cambiare il modo in cui `vettore_selezione` è calcolato.")
test_object("poker_giorni_vincenti",
            incorrect_msg =  "Sembra che `poker_giorni_vincenti` non contenga il risultato corretto. Utilizza `vettore_poker[vettore_selezione]`.")
success_msg("Ottimo lavoro! Continua verso il prossimo esercizio.")
```


--- type:NormalExercise xp:100 skills:1 key:47735366db
## Selezione avanzata

Esattamente come hai appena fatto per il gioco del poker, ti piacerebbe sapere in quali giorni hai vinto alla roulette.

*** =instructions
- Crea la variabile `vettore_selezione`, questa volta però per visualizzare i profitti ottenuti alla roulette nei differenti giorni.
- Assegna i soli risultati positivi alla variabile `roulette_giorni_vincenti`. Questo vettore adesso contiene le vincite del `vettore_roulette`.

*** =hint
Una volta che hai ottenuto il `vettore_selezione`, puoi utilizzare nuovamente `vettore_roulette[vettore_selezione]` per selezionare i risultati positivi da `vettore_roulette`.

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# Vincite al poker e alla rulette da lunedì a venerdì:
vettore_poker <- c(140, -50, 20, -120, 240)
vettore_roulette <- c(-24, -50, 100, -350, 10)
vettore_giorni <- c("Lunedì", "Martedì", "Mercoledì", "Giovedì", "Venerdì")
names(vettore_poker) <- vettore_giorni
names(vettore_roulette) <- vettore_giorni

# In quali giorni hai vinto alla roulette?
vettore_selezione <-

# Seleziona da vettore_roulette questi giorni
roulette_giorni_vincenti <- 
```

*** =solution
```{r}
# Vincite al poker e alla rulette da lunedì a venerdì:
vettore_poker <- c(140, -50, 20, -120, 240)
vettore_roulette <- c(-24, -50, 100, -350, 10)
vettore_giorni <- c("Lunedì", "Martedì", "Mercoledì", "Giovedì", "Venerdì")
names(vettore_poker) <- vettore_giorni
names(vettore_roulette) <- vettore_giorni

# In quali giorni hai vinto alla roulette?
vettore_selezione <- vettore_roulette > 0

# Seleziona da vettore_roulette questi giorni
roulette_giorni_vincenti <- vettore_roulette[vettore_selezione]
```

*** =sct
```{r}
msg = "Non cambiare nulla nella dichiarazione e assegnazione nomi dei vettori `vettore_poker` e `vettore_roulette`."
test_object("vettore_giorni", undefined_msg = msg, incorrect_msg = msg)
test_object("vettore_poker", eq_condition = "equal", undefined_msg = msg, incorrect_msg = msg)
test_object("vettore_roulette", eq_condition = "equal", undefined_msg = msg, incorrect_msg = msg)
test_object("vettore_selezione", 
            incorrect_msg = "Sembra che `vettore_selezione` non contenga il risultato corretto. Utilizza `vettore_roulette > 0`.")
test_object("roulette_giorni_vincenti",
            incorrect_msg = "Sembra che `roulette_giorni_vincenti` non contenga il risultato corretto. Utilizza `vettore_roulette[vettore_selezione]`.")

success_msg("Fantastico! Questo esercizio conclude il capitolo sui vettori. Nel capitolo successivo sarai introdotto alla versione bidimensionale dei vettori: le matrici.")
```



